--- 
title_meta  : Chapter 2
title       : Vetores
description : "Neste curso gratuito de R, nós levaremos você numa viagem a Las Vegas, onde você aprenderá a analisar os resultados das suas apostas utilizando vetores em R! Depois de completar este capítulo, você estará apto a criar vetores em R, nomea-los, selecionar elementos deles e compara-los com vetores diferentes."

--- type:NormalExercise xp:100 skills:1 key:d9b453dbdd
## Crie um vetor

Está se sentindo com sorte? Tomara que sim, porque este capítulo leva você numa viagem à Cidade do Pecado, também conhecida como *Paraíso dos Estatísticos*!

Graças ao R e às suas novas habilidades de análise de dados, você aprenderá como aprimorar sua performance nas mesas de jogo e como incendiar sua carreira de apostador profissional. Este capítulo mostrará como você pode facilmente registrar o progresso de suas apostas e como você pode fazer algumas análises simples sobre as suas ações passadas. Próxima parada, Vegas Baby... VEGAS!!

*** =instructions
- Você ainda lembra do que aprendemos no primeiro capítulo? Atribua o valor `"Go!"` à variável `vegas`. Lembre-se: R é sensível à caixa ALTA ou baixa!

*** =hint
Só digite a seguinte linha no editor:
```
vegas <- "Go!"
```

*** =pre_exercise_code
```{r}
# no pec
```

*** =sample_code
```{r}
# Defina a variável vegas
vegas <- 
```

*** =solution
```{r}
# Defina a variável vegas
vegas <- "Go!"
```

*** =sct
```{r}
test_object("vegas", incorrect_msg = "Certifique-se de que você atribuiu o valor correto a `vegas`. Não esqueça que o R é sensível à caixa ALTA ou baixa!")
success_msg("Ótimo! Vá adiante para o próximo exercício.")
```
